Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ling Lu"
COLLABORATORS = ""
```

---

# Homework 1

Submit your assignment as a Jupyter notebook file called `macid_hw1.ipynb`, where `macid` is replaced with your MacID from your McMaster email address. (For example: since my McMaster email address is pocucan@mcmaster.ca, I would submit a file called `pocucan_hw1.ipynb`.) Submit this file to the appropriate MacDrive link by **OCTOBER 9TH 2020, 11:59 PM **

To complete your assignment, edit the function definitions in this template (then rename the file appropriately, as above). (Don't forget to use appropriate `return()` statements in your functions!) **Do not alter any other part of this template.**

While developing and testing your code, it might be helpful to use
`print` statements, but please remove them before submitting your solution. The template contains several `assert` statements that when executed will test your code.

Any file submitted for grading that does not conform to the specifications above will lead to a grade of zero on the assignment.

Before submitting your solution, you should make sure that it runs properly. To do this, you can open it in the Jupyter notebook server that comes installed with Anaconda, or you can use the McMaster Jupyter server, by going to the website https://mcmaster.syzygy.ca/. (Try running `Kernel > Restart & Run all` from the menu to run all of your code from scratch and make sure nothing is left out.) Feel free to use the Spyder IDE, or some other IDE, to develop your code before entering it into the Jupyter notebook that you submit.

Your grade for each question will depend on whether or not your code correctly handles not just the test cases provided in the template, but other test cases that will be run on your solutions.

Do not leave this assignment until the last minute; last-minute computer/internet issues on your part are your problem, not ours ...

Late assignments will be penalized 25% per day.

All work submitted for grading must be your own. You may discuss homework problems and related material with other students, but you must not submit work copied from others or from the internet.  If you use any internet resources for help, you **must** list them as comments (##) in your answer. 

Do not copy code from any other students, this is the fastest way for me to give you a 0. 

### REMOVE ALL raise NotImplemented() lines from your code!

### REMOVE ALL raise NotImplemented() lines from your code!

```python
import numpy as np # import numpy module 
import matplotlib.pyplot as pyplot # import pyplot module 
```

## 1. Sigmoid Kernels (4 points)

A Kernel function measures the similarity of two arguments. The Sigmoid Kernel was inspired by the multi-layer perceptron. For given ``np.matrix`` vectors ``x`` and ``y``, of size `(1,n)`, write the function `sigmoid_kernel(x,y,gamma,r)`: 

$$ k(x,y) = \tanh \left(\gamma\mathbf{x}\mathbf{y}^T+r \right) $$ .

The function must return a float.

<!-- grade_id: cell-aa041a62adb61ef2, points: 0, grade: false, solution: true -->
```python
def sigmoid_kernel(x,y,gamma,r):
    """Calculates the sigmoid kernel for vectors x and y. both x and y are horizontal vector inputs.
    x : (1,n) np.matrix 
    y : (1,n) np.matrix
    gamma : float parameter
    r : float parameter
    """

    # YOUR CODE HERE
    y_t=np.matrix.transpose(y)
    xy_t=(x*y_t)[0,0]
    return(np.tanh(gamma*xy_t+r))
```

<!-- grade_id: cell-a39104464fc7e94e, points: 4, grade: true, solution: false -->
```python
# tests are hidden here... 

```

## 2. Gaussian Process Kernels (6 points)

GP kernels are often selected to handle the covariance between some vectors. For given ``np.matrix`` vectors ``x`` and ``y``, of size `(1,n)`, write the function `squared_iso(x,y,sigma,l)` that constructs an `(n,n)` matrix using the following form: 

$$ K(x,y) = \left\{  \sigma\exp \left(\frac{-1}{2l^2}(x_i - y_j)^2 \right)   \right\}_{i,j \in 1,\dots, n} $$ 

The function must return an `np.matrix` of size `(n,n)` with each entry `i,j` corresponding to the formula above. Assume all inputs `x,y,sigma,l` are valid inputs. Also note that `i` indexes the entries of the $x$ vector, and `j` the $y$ vector.



<!-- grade_id: cell-b8f45eb736262e13, points: 0, grade: false, solution: true -->
```python
def squared_iso(x,y,sigma,l):
    """Calculates the squared exponential kernel for vectors x and y. x and y are both horizontal vectors. 
    x : (1,n) np.matrix 
    y : (1,n) np.matrix
    sigma : float parameter
    l : float parameter
    """

    # YOUR CODE HERE
    n=x.size
    A=(np.matrix(np.zeros(n*n)).reshape(n,n))
    for i in range(n):
        for j in range(n):
            k= sigma*np.exp(((-1)*(x[0,i]-y[0,j])**2)/(2*(l**2)))
            A[i,j]=k
    return A
        
```

<!-- grade_id: cell-253c286b440ceb3d, points: 6, grade: true, solution: false -->
```python
# tests are hidden here... 

```

### REMOVE ALL raise NotImplemented() lines from your code!

## 3. The Reshaper (6 points)

For a given `np.matrix` variable `x` of size `(1,n)` there may exist a number of possible reshaped matrices. For example, for a matrix `A = np.matrix([1,2,3,4,5,6])` there exists a `(1,6),(2,3),(3,2),(6,1)` possible reshaped matrices. The first one being the original matrix `A`. Write a function `reshape_collector(X)` that takes in a matrix of size `(1,n)` and returns a `list` of all possible reshaped matrices. We will only consider matrices of dimension two. 

For example, an input matrix of ``np.matrix([1,2,3,4])`` will return a list as such:

``
[np.matrix([1,2,3,4]),
np.matrix([[1,2],
           [3,4]]),
np.matrix([[1],
           [2],
           [3],
           [4]])]
``

Note that they must be in a similar order as `(1,6),(2,3),(3,2),(6,1)`.







<!-- grade_id: cell-2e251c573d1c3126, points: 0, grade: false, solution: true -->
```python
def reshape_collector(X):
    """ Takes in a np.matrix X and returns a list of all possible reshaped matrices
    X : (1,n) np.matrix 
    """
    # YOUR CODE HERE
    A=[]
    length=X.size
    for i in range (1,(length+1)):
        if length%i == 0:
            y=np.reshape(X,(i,-1))
            A.append(y)

    return A   
```

<!-- grade_id: cell-0bec16c0ee66e652, points: 6, grade: true, solution: false -->
```python
# tests are hidden here... 



```

## 4. Nomenclature 

For the following questions, these are multi-input multiple choice. There is a variable called ``answer = ""`` which is currently an empty string. Using the document string defined in ``` """<answer options>"""``` please put the appropriate answers. For example, suppose that $S$ is a Non-Autonomous 2nd order system, then for `question_1()`, you would put `answer = "BF"`, the order of letters does not matter. Make sure they are capitalized letters. 

-------------------------------------------------------

```
def question_1():
    """ 
    Select all that apply 
    
    A) Autonomous
    B) Non-Autonomous 
    C) Linear 
    D) Non-Linear
    E) 1st Order
    F) 2nd Order
    """
    
    answer = ""
    
    return(answer) 
 ```
 -----
 
 So I would change ```answer``` to be ```answer = "BF"```.

### 4.0 Practice (0 points)

The following discrete dynamical system $S$ is defined as follows: 

$S$ has no order, is non linear, and is a non-autonomous system. 


Please select all that apply in the `question_0()` function.

<!-- grade_id: cell-ef34e8a62d54a363, points: 0, grade: false, solution: true -->
```python
def question_0():
    """
    Select all that apply 
    A) Autonomous
    B) Non-Autonomous 
    C) Continous
    D) Non-Linear
    E) 1st Order
    F) No Order
    """
    
    answer = "BDF"
    # YOUR CODE HERE
    return(answer)

```

<!-- grade_id: cell-6676e0753ed28540, points: 0, grade: true, solution: false -->
```python
# you must have this cell run, in order to understand how to answer each of the following questions.
assert question_0() == "BDF"
```

### 4.1 (1 point)

The following discrete dynamical system $S$ is defined as follows: 

$$ S = \{ t \in \mathbb{Z}, x^{(t)} = x^{(t-1)} + x^{(t-2)} \} $$


Please select all that apply in the `question_1()` function.

<!-- grade_id: cell-dac36d38b0be4afa, points: 0, grade: false, solution: true -->
```python
def question_1():
    """
    Select all that apply 
    
    A) Autonomous
    B) Non-Autonomous 
    C) Linear 
    D) Non-Linear
    E) 1st Order
    F) 2nd Order
    """
    
    answer = "ACF"
    # YOUR CODE HERE
    
    
    return(answer) 

```

<!-- grade_id: cell-3bdd6fc2edc9e429, points: 1, grade: true, solution: false -->
```python
# hidden tests beware... 
```

### 4.2 (1 point)

The following discrete dynamical system $S$ is defined as follows: 

$$ S = \{ t \in \mathbb{Z}, x^{(t)} = x^{(t-1)}x^{(t-2)} + \sin(2\pi t + \frac{\pi}{2})\} $$


Please select all that apply in the `question_2()` function.


<!-- grade_id: cell-1eeefe3c5fbdd7f0, points: 0, grade: false, solution: true -->
```python
def question_2():
    """
    Select all that apply 
    
    A) Autonomous
    B) Non-Autonomous 
    C) Linear 
    D) Non-Linear
    E) 1st Order
    F) 2nd Order
    """
    
    answer = "BDF"
    # YOUR CODE HERE
    
    
    return(answer) 

```

<!-- grade_id: cell-23976b1932e01bb7, points: 1, grade: true, solution: false -->
```python
# hidden tests beware... 
```

### 4.3 (1 point)

The following discrete dynamical system $S$ is defined as follows: 

$$ S = \{ t \in \mathbb{Z}, x^{(t)} = 5x^{(t-1)} + 2 \} $$


Please select all that apply in the `question_3()` function.


<!-- grade_id: cell-8d3995117d74b25e, points: 0, grade: false, solution: true -->
```python
def question_3():
    """
    Select all that apply 
    
    A) Autonomous
    B) Non-Autonomous 
    C) Linear 
    D) Non-Linear
    E) 1st Order
    F) 2nd Order
    """
    
    answer = "ACE"
    # YOUR CODE HERE
   
    
    return(answer) 

```

### 4.4 (1 point)

The following discrete dynamical system $S$ is defined as follows: 

$$ S = \{ t \in \mathbb{Z}, x^{(t)} = 5tx^{(t-2)} + \cos(2\pi t)\} $$


Please select all that apply in the `question_4()` function.


<!-- grade_id: cell-9e21c9914afa134b, points: 0, grade: false, solution: true -->
```python
def question_4():
    """
    Select all that apply 
    
    A) Autonomous
    B) Non-Autonomous 
    C) Linear 
    D) Non-Linear
    E) 1st Order
    F) 2nd Order
    """
    
    answer = "BDF"
    # YOUR CODE HERE
    
    
    return(answer) 

```

<!-- grade_id: cell-c2a554829bb9bd7e, points: 1, grade: true, solution: false -->
```python

```

### First Order Discrete Dynamical Systems 

The following class `FirstOrder_DDS` encompases all of first order discrete dynamical system problems. The following questions will require you to program a function for each problem, and then pass them into the class along with an initial condition `ic`. These problems are similar to the topics we covered in class.


```python
class FirstOrder_DDS:
    """ Do not change these methods."""
    
    def __init__(self,f,ic):
        self.f = f
        self.history = np.zeros(200)
        self.history[0] = ic
        self.ic = ic 
        self.step = 0
        self.current = ic 


    def step_forward(self):
        
        if self.step == len(self.history) - 1:
            self.history = np.append(self.history,np.zeros(200))
            
        old = self.history[self.step]
        new = self.f(old)
        self.step += 1
        self.history[self.step] = new
        self.current = new
        
        

    def plot(self,xlab = "Time",ylab = "Output"):
        time_steps = range(0,self.step)
        output = self.history[ :self.step ]
        pyplot.plot(time_steps,output)
        pyplot.xlabel(xlab)
        pyplot.ylabel(ylab)

```

### 5 Chloe's Cattery

Chloe owns a cattery, she currently owns two cats Bella and Nova who are considered to be the first breeding cats. In addition, she also owns a breeder named Milo who will mate with each cat. 

Every year, Chloe allows one litter for each breeding cat, who for the purposes of this problem, have exactly 5 kittens per litter. After each year, she sells both of the breeding cats, and retains the kittens. She also sells off the breeder.
 
For the next year, she selects two kittens to act as female breeders and buys a new male breeder. 

Let $c^{(t)}$ denote the number of cats at year $t$. 



### Chloe's Cattery 5.1 (1 point)


Which of the following difference equations will model the number of cats in Chloe's cattery.  
$$$$

$A:\quad c^{(t)} =  2c^{(t-1)} + 1 $

$B:\quad c^{(t)} =  c^{(t-1)} + 8 $

$C:\quad c^{(t)} =  c^{(t-1)} - 3 $

$D:\quad c^{(t)} =  2c^{(t-1)} - 2 $


Please select an answer in question `question_51` by writing in the string `answer = ""`. For example if you think the answer is `E` then write `answer = "E"`. Hint... the answer is not E.

<!-- grade_id: cell-d778830f21e979ee, points: 0, grade: false, solution: true -->
```python
def question_51():
    """Change answer based on the question above. """
    
    answer = "B"
    # YOUR CODE HERE
    
    return(answer)
```

<!-- grade_id: cell-6954bba5b967c9bb, points: 1, grade: true, solution: false -->
```python
# hidden tests beware...
```

### Chloe's Cattery 5.2 (1 point)
Using the selected equation in 5.1, code up the function `breed_cats(c)` which takes in an argument `c` for the number of cats and returns the new population of cats.

<!-- grade_id: cell-0ffc23ebef185e0e, points: 0, grade: false, solution: true -->
```python
def breed_cats(c):
    """
    Breeds two cats, and have litter.
    """
    new_value = 0
    # YOUR CODE HERE
    new_value = c+8
    return(new_value)
```

<!-- grade_id: cell-87b4bfe95243a44c, points: 1, grade: true, solution: false -->
```python
# hidden tests beware...
```

### Chloe's Cattery 5.3 (1 point)

Given an intial population of 3 cats. How years will it take for chloe to have 27 cats or greater? Use the `FirstOrder_DDS` class along with the appropriate methods to answer `question_53`. `answer` is an integer number, and is currently set to 0. Remember, the current year counts as 0.

I suggest you do your rough work in another notebook by copying the class definition over there `FirstOrder_DDS` and the function `breed_cats`. 

Remember, I am looking for a single integer number for the `answer`.

<!-- grade_id: cell-c728e22c0e481839, points: 0, grade: false, solution: true -->
```python
def question_53():
    """
    How many years will it take to have 27 cats or greater?
    """
    answer = 0
    # YOUR CODE HERE
    year_need=FirstOrder_DDS(breed_cats,3)
    for i in range(1,100):
        year_need.step_forward()
        if year_need.current < 27:
            answer = year_need.step+1
    
    return(answer)
    
```

<!-- grade_id: cell-00bb45b5adf507c7, points: 1, grade: true, solution: false -->
```python
# hidden tests beware
```

### Chloe's Cattery 5.4 (2 point)
 
Given an initial population of 8 cats. Suppose Chloe runs her breeding business indefinity, and wants to know how much money she need to spend on cat food. Use the `FirstOrder_DDS` class along with the appropriate methods to create a function called `cat_tracker` to keep track of how many cats she will have by year.

The `cat_tracker` function should take a `year` and return the number of cats she will have by year. The function returns `answer` as a `list` type.

Hint.. you should use the `history` value in the class...

<!-- grade_id: cell-eda66ba4ebc12524, points: 0, grade: false, solution: true -->
```python
def cat_tracker(year):
    """
    Collects the number of cats by year. 
    """
    
    answer = 0
    # YOUR CODE HERE
    tracker=FirstOrder_DDS(breed_cats,8)
    for i in range (1,year+1):
        tracker.step_forward()
        if tracker.history[i] != 0:
            y = tracker.history[i]
            answer.append(y)
    return(answer)

```

<!-- grade_id: cell-e25559b63df602c2, points: 2, grade: true, solution: false -->
```python
# hidden tests beware
```

### 6 Jason's RRSP
Jason works as a Data Scientist lead. He makes a whooping of 200k a year, from which 40% of it gets snatched by the CRA. The rest of it he spends on rent (30k a year) and other living expenses (50k a year). Jason saves the reminder in his RRSP account which grows at 5% a year.

Let r(ð‘¡) denote the value of his RRSP account at time t.

### Jason's RRSP 6.1 (1 point)


Which of the following difference equations will model the value of Jason's RRSP account (in thousands).  
$$$$

$A:\quad r^{(t)} =  r^{(t-1)} + 40 $

$B:\quad r^{(t)} =  1.05r^{(t-1)} + 120 $

$C:\quad r^{(t)} =  1.05r^{(t-1)} + 40 $

$D:\quad r^{(t)} =  r^{(t-1)} + 120 $


Please select an answer in question `question_61` by writing in the string `answer = ""`. For example if you think the answer is `E` then write `answer = "E"`. Hint... the answer is not E.

<!-- grade_id: cell-dd8a4f41a9d139e4, points: 0, grade: false, solution: true -->
```python
def question_61():
    """Change answer based on the question above. """
    
    answer = "C"
    # YOUR CODE HERE
   
    return(answer)
```

<!-- grade_id: cell-fc40b54dd9cb13dc, points: 1, grade: true, solution: false -->
```python
# hidden tests beware...
```

### Jason's RRSP 6.2 (1 point)
Use your equation choice from 6.1, and create a `rrsp_value(v)` function which takes in an argument v for current value and returns the new RRSP account value for next year.

<!-- grade_id: cell-0ae6d9b3a1d939e2, points: 0, grade: false, solution: true -->
```python
def rrsp_value(v):
    """
    get jasons rrsp
    """
    new_value = 0
    # YOUR CODE HERE
    new_value=v*1.05+40
    return(new_value)
```

<!-- grade_id: cell-8a10e417e72e90b2, points: 1, grade: true, solution: false -->
```python
# hidden tests beware...
```

### Jason's RRSP 6.3 (2 point)

Jason has 400k saved up and he wants to buy a house. As a first time home buyer he can use the money from his RRSP. How many years will it take him to have enough for a 1 millon dollar home?

Remember, I am looking for a single integer number for the `answer`.

<!-- grade_id: cell-8105f0932f3f47cd, points: 0, grade: false, solution: true -->
```python
def question_63():
    """
    How many years will it take to have 1000k or great?
    """
    answer = 0
    # YOUR CODE HERE
    ic=400
    years=FirstOrder_DDS(rrsp_value,ic)
    for i in range(0,100):
        years.step_forward()
        if years.current <1000:
            answer=years.step+1
    return(answer)
```

<!-- grade_id: cell-f4e49a2763581151, points: 2, grade: true, solution: false -->
```python
# hidden tests beware...
```

### Jason's RRSP 6.4 (2 point)
You work as a financial advisor at the bank. Jason wants you to give a value projection of his RRSP account (Currently, he has 100k in the account). Use the FirstOrder_DDS class along with the appropriate methods to create a function called `value_projection` to keep track of how Jason's account balance over the years.

The `value_projection` function should take a year and return all of the account value projections up to that year. The function must returns `answer` as a List.

(hint this is like the jasons history over time, except you are returning a list instead of np.array)

<!-- grade_id: cell-f7696d5e79a50c78, points: 0, grade: false, solution: true -->
```python
def value_projection(year):
    """
    Collects the account value by year. 
    """
    
    answer = 0
    # YOUR CODE HERE
    value=FirstOrder_DDS(rrsp_value,100)
    for i in range (1,year+1):
        value.step_forward()
        if value.history[i] != 0:
            answer.append(value.history[i])
    return(answer)
```

<!-- grade_id: cell-8fd0c4f5ee1dc3c8, points: 2, grade: true, solution: false -->
```python
# hidden tests beware
```

## Multi-State Dynamical Systems (Includes Higher Order)

For the following questions, you can use the `MultiDDS` class definition in the next cell to answer the questions. 




```python
class MultiDDS:

    def __init__(self,f,ic):
        # ic corresponds to the total number of states, but first check to make sure its a matrix. 
        if type(ic) != np.matrix:
            raise Exception("initial conditions not of np.matrix type")

        if ic.shape[0] == 1:
            raise Exception("inital conditions need to have vertical matrix vectors")


        p = ic.shape[0]
        # set inital conditions
        self.ic = ic 
        self.p = p 

        # set history and other variables
        self.history = np.matrix(np.zeros((p,int(1e4))))
        self.history[:,0] = ic 
        self.f = f 
        self.step = 0
        self.current = ic


    def step_forward(self):
        new = self.f(self.history[:,self.step])
        self.step += 1 
        self.history[:,self.step] = new 
        self.current = new 


    def plot(self,position = 0,xlab = "Input",ylab = "Output"):
        steps = range(0,self.step)
        pyplot.plot(steps, self.history[position,:self.step].transpose())
        pyplot.xlabel(xlab)
        pyplot.ylabel(ylab)

```

### 7.1 Wolf/Rabbit Populations (2 points)

The wolf and rabbit populations measured in thousands ($w$ and $r$) are modelled for year `t` with the following equations: 

$$w^{(t)} = -0.03w^{(t-1)} + 0.06r^{(t-1)}$$
$$ r^{(t)} = -0.06w^{(t-1)} + 1.03r^{(t-1)} $$


Write the function `animal_bogaloo` that takes in a vector `x` of shape `(p,1)` increments the states one year forward according to the difference equations above, and returns the new state vector. Must return a `(p,1)` matrix.  Remember, `p` is the dimension of your state vector `x`. 

<!-- grade_id: cell-095646072fa9fdf2, points: 0, grade: false, solution: true -->
```python
def animal_bogaloo(x):
    """Wolf eat rabbit, rabbit eat grass"""
    
    # YOUR CODE HERE
    A=np.matrix([[-0.03,0.06],[-0.06,1.03]])
    return (A*(x.reshape(2,1)))
```

<!-- grade_id: cell-549c16983022432a, points: 2, grade: true, solution: false -->
```python
# hidden tests beware
```

### 7.2 Modelling Bugs Bunny (2 points)
Given the intial conditions: 10004 wolves, and 20005 rabbits what is the population of rabbits after 7 years? You can create a new notebook for rough work and use the predefined class above. I am looking for a single floating point number in the `answer` variable. `answer` is currently set to 0. Round to the nearest 3 decimal places AFTER the final calculations. 

<!-- grade_id: cell-d826eebb7bfff39e, points: 0, grade: false, solution: true -->
```python
def question_animal_1():
    """
    What is the population of rabbits after 7 years using the predefined initial conditions. 
    """
    
    answer = 0
    # YOUR CODE HERE
    ic=np.matrix([10.004,20.005]).reshape(2,1)
    pop_of_rabbits=MultiDDS(animal_bogaloo,ic)
    actual_value=0
    for i in range(0,100):
        pop_of_rabbits.step_forward()
        if i == 6:
            actual_value=pop_of_rabbits.current
            answer=np.round(actual_value[1,0],3)
    return(answer)
    
```

<!-- grade_id: cell-bab3a8a40101485d, points: 2, grade: true, solution: false -->
```python
# hidden tests beware...
```

### 7.3 Wolf Adjustments  ( 2 points)

The wolf and rabbit populations in thousands ($w$ and $r$) are modelled for year `t`.  Unfortuntely, the rabbit populations have become out of control. As a reduction measure, the Parks department of Pawnee has decided to transfer 25 new wolves into the wolf population each year. Using the previously defined equations, modify the wolf difference equation accordingly. 

Once that has been completed, write the function `animal_bogaloo_2` that takes in a vector `x` of shape `(p,1)` increments the states one year forward according to the difference equations above, and returns the new state vector. Must return a `(p,1)` matrix. Note that `p` may differ from the previous question. 

<!-- grade_id: cell-d7cb373c7ce92ac1, points: 0, grade: false, solution: true -->
```python
def animal_bogaloo_2(x):
    """Wolf eat rabbit, rabbit eat grass, too many rabbits, introduce wolves"""
    
    # YOUR CODE HERE
    A_new=np.matrix([[-0.03,0.06],[-0.06,1.03]])
    b=np.matrix([25,0]).reshape(2,1)
    return (A_new*x+b)
```

<!-- grade_id: cell-4611727b9150a5a5, points: 2, grade: true, solution: false -->
```python
# beware of hidden tests 
```

### 8 Global Economy and Geopolitics

The world GDP is measure in trillons. Consider a discrete dynamics model of the top 5 countries in the world by GDP ranking (U = USA, C = China, J = Japan, G = Germany and K = UK): 

$$ U^{(t)} = 1.01U^{(t-1)} + 0.02C^{(t-1)} + 0.01J^{(t-1)} + 0.01G^{(t-1)} + 0.01K^{(t-1)} + 0.01$$
$$ C^{(t)} = 0.03U^{(t-1)} + 1.04C^{(t-1)} + 0.01J^{(t-1)} + 0.01G^{(t-1)} - 0.01K^{(t-1)}  + 0.03$$
$$ J^{(t)} = 0.01U^{(t-1)} + 0.02C^{(t-1)} + 1.00J^{(t-1)} - 0.03G^{(t-1)} + 0.01K^{(t-1)}  - 0.02$$
$$ G^{(t)} = 0.03U^{(t-1)} + 0.03C^{(t-1)} - 0.03J^{(t-1)} + 1.02G^{(t-1)} + 0.03K^{(t-1)}  + 0.02$$
$$ K^{(t)} = 0.01U^{(t-1)} + 0.02C^{(t-1)} + 0.01J^{(t-1)} + 0.02G^{(t-1)} + 1.01K^{(t-1)}  + 0.01$$. 

### GDP from Trading 8.1 (1 point)


With out the effect of global trading, which country would experience a decrease of GDP.  
$$$$

$A: USA$

$B: China$

$C: Japan$

$D: Germany$

$D: UK$


Please select an answer in question `question_81` by writing in the string `answer = ""`. For example if you think the answer is `E` then write `answer = "E"`. Hint... the answer might or might not E.

<!-- grade_id: cell-86a3e94a7f96d249, points: 0, grade: false, solution: true -->
```python
def question_81():
    """Change answer based on the question above. """
    
    answer = "C"
    # YOUR CODE HERE
    
    return(answer)
```

<!-- grade_id: cell-271dfeb34f8d7817, points: 1, grade: true, solution: false -->
```python
# hidden tests beware
```

### GDP from Trade 8.2 ( 2 points)
Write the function `gdp` that takes in a vector x of shape (5,1) increments the states one year forward according to the difference equations above, and returns the new state vector. Must return a (5,1) matrix. (Note p = 5, and is given to you... this means you do not have to add a 1 to the end of the matrix to handle coefficients) Instead, think about x = Ax + b form...

<!-- grade_id: cell-b884ba2e513f854b, points: 0, grade: false, solution: true -->
```python
def gdp(x):
    """How countries effect each other's GDP"""
    # YOUR CODE HERE
    A=np.matrix([[1.01,0.02,0.01,0.01,0.01],
                 [0.03,1.04,0.01,0.01,-0.01],
                 [0.01,0.02,1.00,-0.03,0.01],
                [0.03,0.03,-0.03,1.02,0.03],
                [0.01,0.02,0.01,0.02,1.01]])
    b=np.matrix([0.01,0.03,-0.02,0.02,0.01])
    gdp=A*x.transpose()+b.transpose()
    return (gdp)
```

<!-- grade_id: cell-d89dec1a6b3a2253, points: 2, grade: true, solution: false -->
```python
# hidden tests beware
```

### Trade War! Tariff Effects 7.3  ( 3 points)

One country decided to impose tariff, this caused a domino effect in which in every country imposed tariff on each other. The result of this is a 300 basis point (0.03 trillion) shift downward across the GDP transition matrix.

Apply this change to the `gdp` function, call the new function `gdp_tariff`. 

<!-- grade_id: cell-3aedfbb487fa37da, points: 0, grade: false, solution: true -->
```python
def gdp_tariff(x):
    """How countries effect each other's GDP"""
    # YOUR CODE HERE
    A_83=np.matrix([[1.01,0.02,0.01,0.01,0.01],
                 [0.03,1.04,0.01,0.01,-0.01],
                 [0.01,0.02,1.00,-0.03,0.01],
                [0.03,0.03,-0.03,1.02,0.03],
                [0.01,0.02,0.01,0.02,1.01]])
    b=np.matrix([0.01,0.03,-0.02,0.02,0.01])
    T=np.full((5,5),-0.03)+np.diag([0.03,0.03,0.03,0.03,0.03])
    A_tariff=A_83+T
    gdp_adj=A_tariff*x.transpose()+b.transpose()
    return (gdp_adj)
```

<!-- grade_id: cell-dfd119e294bb53f9, points: 3, grade: true, solution: false -->
```python
# hidden tests beware
```

# THIS IS THE END OF THE ASSIGNMENT, PLEASE SUBMIT TO THIS MACDRIVE LINK WITH THIS ASSIGNMENT RENAMED ACCORDING TO YOUR MACID 

https://macdrive.mcmaster.ca/u/d/818705de222f47838a63/
